<a href="https://colab.research.google.com/github/syalam1998/hds5210-2023/blob/main/week05/week05_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
def longest(words):

  # I am including doctest example for this function
    """
    (String) -> int
    This function finds out longest word from the given sentence

    This version expects a sentence

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    """

    longest_word = words[0]
    longest_position = 0

    for i, word in enumerate(words):
        if len(word) > len(longest_word):
            longest_word = word
            longest_position = i

    return longest_position

if __name__ == "__main__":
    import doctest
    doctest.testmod()
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
def f_to_c(tempf):

  # I am including doctest with tests for 212f, 32f, and 98.6f
    """
    (float) -> float
    This function converts the farenheit by given farenheit to celsius conversion formula

    This Function takes input as float and returns celsius as float

    >>> f_to_c(212)
    100.0

    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0

    """

    tempc = 5/9 * (tempf - 32)
    return tempc

if __name__ == "__main__":
    import doctest
    doctest.testmod()
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import datetime

def los(admit_date, discharge_date, diagnosis):

    """
    (date, date, String) -> int

    This Function calculates the length of stay, in whole days, between the admit date and discharge date

    if the diagnosis is "Observation" then the length will be returned as 0 regardless of the admit and discharge dates.
    This function takes three parameters: a start date, an end date, and a diagnosis.

    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5

    >>> los(date(2019,1,5),  date(2019,1,9), "Hypertension")
    4

    >>> los(date(2019,1,12), date(2019,1,13), "Anxiety")
    1

    >>> los(date(2019,1,15), date(2019,1,16), "Observation")
    0

    >>> los(date(2019,1,4), date(2019,1,4), "Anxiety")
    0
    """
    if diagnosis == "Observation":
        return 0
    else:
    # Calculate the length of stay in days
      return (discharge_date - admit_date).days


if __name__ == "__main__":
    import doctest
    doctest.testmod()
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2019,1,5),  date(2019,1,9), "Hypertension")
Expecting:
    4
ok
Trying:
    los(date(2019,1,12), date(2019,1,13), "Anxiety")
Expecting:
    1
ok
Trying:
    los(date(2019,1,15), date(2019,1,16), "Observation")
Expecting:
    0
ok
Trying:
    los(date(2019,1,4), date(2019,1,4), "Anxiety")
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

def average_los(encounters):

    """
    (dictionary) -> float

    This Function calculates the average of the LOS for the encounters in the provided dictionary.
    It takes input as a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.

    If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

    >>> dictionary = {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },"E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },"E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },"E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },"E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}}
    >>> round(average_los(encounters),2)
    3.33

    """
    total_los = 0
    count = 0

    # Iterate through each encounter in the dictionary
    for encounter in encounters.values():
        if 'admit' in encounter and 'discharge' in encounter and 'diagnosis' in encounter:
            admit_date = encounter['admit']
            discharge_date = encounter['discharge']
            diagnosis = encounter['diagnosis']

            # Calculate the length of stay for the encounter using the los function
            los_value = los(admit_date, discharge_date, diagnosis)

            # Only consider encounters with a non-zero LOS
            if los_value > 0:
                total_los += los_value
                count += 1

    if count == 0:
        return 0
    else:
        return total_los / count

if __name__ == "__main__":
    import doctest
    doctest.testmod()
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    dictionary = {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },"E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },"E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },"E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },"E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}}
Expecting nothing
ok
Trying:
    round(average_los(encounters),2)
Expecting:
    3.33
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
def c_to_f(tempc):
  """
  (float) -> float

  This Function converts degrees Celsius to degrees Farenheit

  >>> c_to_f(100)
  212.0

  >>> c_to_f(0)
  32.0

  >>> c_to_f(37)
  98.6
  """
  tempf = (tempc * 9/5) + 32
  return tempf

if __name__ == "__main__":
    import doctest
    doctest.testmod()
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.